# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 20% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np
from IPython.display import display



# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    Energy = pd.read_excel("assets/Energy Indicators.xls",usecols='C:F',
                           skiprows=[1,2,3,4,5,6,7,8,9,10,12,13,14,15,16,17],
                          skipfooter=18,
                          names=['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable'])
    
    Energy["Country"] = Energy["Country"].replace("Republic of Korea", "South Korea")
    Energy["Country"] = Energy["Country"].replace("United States of America", "United States")
    Energy["Country"] = Energy["Country"].replace("United Kingdom of Great Britain and Northern Ireland", "United Kingdom")
    Energy["Country"] = Energy["Country"].replace("China, Hong Kong Special Administrative Region", "Hong Kong")
    
    Energy["Energy Supply"] = Energy["Energy Supply"].replace("...",np.NaN)
    Energy["Energy Supply"] = Energy["Energy Supply per Capita"].replace("...",np.NaN)
    Energy["Energy Supply"] = Energy["% Renewable"].replace("...",np.NaN)
    Energy['Country'] = Energy["Country"].replace(r"\(.*?\)", "",regex=True)
    #display(Energy.loc[Energy['Country']=='United States Virgin Islands'])
    #print(Energy.head(50))
    #print(Energy.loc[Energy["Country"]=='South Korea'])
    # YOUR CODE HERE
    GDP = pd.read_csv('assets/world_bank.csv')
    GDP.columns = GDP.iloc[3]
    GDP = GDP.drop(index=[0,1,2,3], axis=0)
    GDP = GDP.reset_index()
    
    #print(GDP.keys())
    GDP["Country Name"] = GDP["Country Name"].replace("Korea, Rep.", "South Korea")
    GDP["Country Name"] = GDP["Country Name"].replace("United States of America", "United States")
    GDP["Country Name"] = GDP["Country Name"].replace("United Kingdom of Great Britain and Northern Ireland", "United Kingdom")
    #display(GDP.keys)
    
    GDP = GDP.rename(columns={"Country Name":"Country"})
    ScimEn = pd.read_excel("assets/scimagojr-3.xlsx")
    
    last10 = GDP[GDP.columns[-10:]]
    last10['Country'] = GDP['Country']
    #display(ScimEn.keys())
    #display(Energy.keys())
    #display(GDP.keys())
    
    Top15 = ScimEn[ScimEn['Rank'] < 16]
    
    #display(Top15)
    #display(last10)
    #display(ScimEn.head())
    #display(Energy.head())
    #display(GDP.head())
    
    TopLat = pd.merge(Top15,last10,how='inner',on='Country')
    TopLat = TopLat.set_index('Country')
    final = pd.merge(TopLat, Energy,how='left',on='Country')
    print('Top15')
    #final = final.set_index('Country')
   # final = 
    display(final)
    #display(final)
    return final
    #display(TopLat)
    #print(Energy.head(-30))
    #raise NotImplementedError()
   
    
answer_one()
```

Output:
```
Top15
               Country  Rank  Documents  Citable documents  Citations  \
0                China     1     127050             126767     597237   
1        United States     2      96661              94747     792274   
2                Japan     3      30504              30287     223024   
3       United Kingdom     4      20944              20357     206091   
4   Russian Federation     5      18534              18301      34266   
5               Canada     6      17899              17620     215003   
6              Germany     7      17027              16831     140566   
7                India     8      15005              14841     128763   
8               France     9      13153              12973     130632   
9          South Korea    10      11983              11923     114675   
10               Italy    11      10964              10794     111850   
11               Spain    12       9428               9330     123336   
12           Australia    14       8831      
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-6-e4f26ba1c257> in <module>
----> 1 assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"
      2 
      3 assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

AssertionError: Q1: You should return a DataFrame!
```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    # YOUR CODE HERE
    GDP = pd.read_csv('assets/world_bank.csv')
    print(len(GDP)-15)
    
    #raise NotImplementedError()
answer_two()
```

Output:
```
253
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    top15 = answer_one()
    display(top15)
    last10 = top15[top15.columns[8:18]]
    mean = last10.mean(axis=1)
    #raise NotImplementedError()
    top15['avgGDP'] = mean
    display(mean)
    return top15
answer_three()
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0      1               China     127050             126767     597237   
1      2       United States      96661              94747     792274   
2      3               Japan      30504              30287     223024   
3      4      United Kingdom      20944              20357     206091   
4      5  Russian Federation      18534              18301      34266   
5      6              Canada      17899              17620     215003   
6      7             Germany      17027              16831     140566   
7      8               India      15005              14841     128763   
8      9              France      13153              12973     130632   
9     10         South Korea      11983              11923     114675   
10    11               Italy      10964              10794     111850   
11    12               Spain       9428               9330     123336   
12    13                Iran       8896               8
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    top15 = answer_one()
    six = top15.loc[5]
    
    #display(six)
    #display(six[2006.0])
    #print(type(six[2006.0]))
    dif = six[2010.0]-six[2006.0]
    #print('Here')
    #print(dif)
    #raise NotImplementedError()
    return dif
answer_four()
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0      1               China     127050             126767     597237   
1      2       United States      96661              94747     792274   
2      3               Japan      30504              30287     223024   
3      4      United Kingdom      20944              20357     206091   
4      5  Russian Federation      18534              18301      34266   
5      6              Canada      17899              17620     215003   
6      7             Germany      17027              16831     140566   
7      8               India      15005              14841     128763   
8      9              France      13153              12973     130632   
9     10         South Korea      11983              11923     114675   
10    11               Italy      10964              10794     111850   
11    12               Spain       9428               9330     123336   
12    13                Iran       8896               8
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    Energy = pd.read_excel("assets/Energy Indicators.xls",usecols='C:F',
                           skiprows=[1,2,3,4,5,6,7,8,9,10,12,13,14,15,16,17],
                          skipfooter=18,
                          names=['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable'])
    
    Energy["Country"] = Energy["Country"].replace("Republic of Korea", "South Korea")
    Energy["Country"] = Energy["Country"].replace("United States of America", "United States")
    Energy["Country"] = Energy["Country"].replace("United Kingdom of Great Britain and Northern Ireland", "United Kingdom")
    Energy["Country"] = Energy["Country"].replace("China, Hong Kong Special Administrative Region", "Hong Kong")
    
    #Energy["Energy Supply"] = Energy["Energy Supply"].replace("...",np.NaN)
    #Energy["Energy Supply"] = Energy["Energy Supply per Capita"].replace("...",np.NaN)
    #Energy["Energy Supply"] = Energy["% Renewable"].replace("...",np.NaN)
    #Energy['Country'] = Energy["Country"].replace(r"\(.*?\)", "",regex=True)
    #mean = Energy['Energy Supply per Capita'].nanmean()
    #mean = Energy['Energy Supply per Capita'].where(type(Energy['Energy Supply per Capita'])=='str')
    #for val in Energy['Energy Supply per Capita']:
     #   print(type(val))
        #rint(i)
    #print(mean)
    mean = pd.to_numeric(Energy['Energy Supply per Capita'],errors='coerce').mean()
    print(mean)
answer_five()
```

Output:
```
90.80269058295964
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    Energy = pd.read_excel("assets/Energy Indicators.xls",usecols='C:F',
                           skiprows=[1,2,3,4,5,6,7,8,9,10,12,13,14,15,16,17],
                          skipfooter=18,
                          names=['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable'])
    
    maxs = Energy.where(Energy['% Renewable'] == Energy['% Renewable'].max()).dropna()
    display(Energy[Energy['% Renewable'] == Energy['% Renewable'].max()])
    print(maxs)
answer_six()
```

Output:
```
Country Energy Supply Energy Supply per Capita  % Renewable
2    Albania           102                       35        100.0
24    Bhutan            63                       83        100.0
117  Lesotho            48                       23        100.0
     Country Energy Supply Energy Supply per Capita  % Renewable
2    Albania           102                       35        100.0
24    Bhutan            63                       83        100.0
117  Lesotho            48                       23        100.0
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    ScimEn = pd.read_excel("assets/scimagojr-3.xlsx")
    print(ScimEn)
    ratio = ScimEn['Self-citations']/ScimEn['Citations']
    #print(ScimEn.where(ScimEn['Citations']==np.NaN))
    print(ratio)
    ScimEn['ratio'] = ratio * 100
    display(ScimEn)
    #raise NotImplementedError()
answer_seven()
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0       1               China     127050             126767     597237   
1       2       United States      96661              94747     792274   
2       3               Japan      30504              30287     223024   
3       4      United Kingdom      20944              20357     206091   
4       5  Russian Federation      18534              18301      34266   
..    ...                 ...        ...                ...        ...   
186   187              Guyana          1                  1          0   
187   188    Christmas Island          1                  1          0   
188   189             Reunion          1                  1          2   
189   190         Saint Lucia          1                  1          0   
190   191          Mauritania          1                  1          1   

     Self-citations  Citations per document  H index  
0            411683                    4.70      138  
1     
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    one = answer_one()
    display(one)
    cit_per_pers = one['Citations']/one['Citations per document']
    
    one['citable docs per person'] = cit_per_pers
    cor = one[['citable docs per person','Energy Supply per Capita']].corr(method='pearson')
    print('cor')
    display(cor)
    display(cit_per_pers)
    #raise NotImplementedError()
answer_nine()
plot9()
```

Output:
```
Top15
               Country  Rank  Documents  Citable documents  Citations  \
0                China     1     127050             126767     597237   
1        United States     2      96661              94747     792274   
2                Japan     3      30504              30287     223024   
3       United Kingdom     4      20944              20357     206091   
4   Russian Federation     5      18534              18301      34266   
5               Canada     6      17899              17620     215003   
6              Germany     7      17027              16831     140566   
7                India     8      15005              14841     128763   
8               France     9      13153              12973     130632   
9          South Korea    10      11983              11923     114675   
10               Italy    11      10964              10794     111850   
11               Spain    12       9428               9330     123336   
12           Australia    14       8831      
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    final = answer_one()
    median = final['% Renewable'].median()
    print(median)
    #for val in final:
     #   print(val['% Renewable'])
    HR = final['% Renewable'] > median
    HR = HR = HR.astype(int)
    final['HighRenew'] = HR
    print(final)
    #display(final)
    #raise NotImplementedError()
    return final
answer_ten()
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0      1               China     127050             126767     597237   
1      2       United States      96661              94747     792274   
2      3               Japan      30504              30287     223024   
3      4      United Kingdom      20944              20357     206091   
4      5  Russian Federation      18534              18301      34266   
5      6              Canada      17899              17620     215003   
6      7             Germany      17027              16831     140566   
7      8               India      15005              14841     128763   
8      9              France      13153              12973     130632   
9     10         South Korea      11983              11923     114675   
10    11               Italy      10964              10794     111850   
11    12               Spain       9428               9330     123336   
12    13                Iran       8896               8
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  #'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    
    final = answer_one()
    stats = pd.DataFrame(columns =['Continent','size', 'sum', 'mean', 'std'] ) 
    stats.set_index('Continent')
    final=final.set_index('Country')
    final['Est Pop'] = final['Energy Supply'] / final['Energy Supply per Capita']
    #stats = final.set_index('Country')
    #display(final)
    #final.set_index(ContinentDict.values())
    for key in ContinentDict:
        final.at[key,'Continent'] = ContinentDict[key]
    #final.agg({"Continent":ContinentDict})
    display(stats)
    #for Continent in  final['Continent'].unique():
    #    print(Continent)
        
    for group,frame in final.groupby(by='Continent'):
        stats.at[group,'sum'] = frame['Est Pop'].sum()
        #repeat this for the std, mean, and len
    display(stats)   
    #raise NotImplementedError()
answer_eleven()
```

Output:
```
Top15
               Country  Rank  Documents  Citable documents  Citations  \
0                China     1     127050             126767     597237   
1        United States     2      96661              94747     792274   
2                Japan     3      30504              30287     223024   
3       United Kingdom     4      20944              20357     206091   
4   Russian Federation     5      18534              18301      34266   
5               Canada     6      17899              17620     215003   
6              Germany     7      17027              16831     140566   
7                India     8      15005              14841     128763   
8               France     9      13153              12973     130632   
9          South Korea    10      11983              11923     114675   
10               Italy    11      10964              10794     111850   
11               Spain    12       9428               9330     123336   
12           Australia    14       8831      
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  #'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    
    final = answer_one()
    
    final["bins"] = pd.cut(final['% Renewable'],5)
    for key in ContinentDict:
        final.at[key,'Continent'] = ContinentDict[key]
    x=final.groupby(['Continent','bins']).size()
    display(x.head(100))
    #display(final)
    #raise NotImplementedError()
answer_twelve()

```

Output:
```
Top15
               Country  Rank  Documents  Citable documents  Citations  \
0                China     1     127050             126767     597237   
1        United States     2      96661              94747     792274   
2                Japan     3      30504              30287     223024   
3       United Kingdom     4      20944              20357     206091   
4   Russian Federation     5      18534              18301      34266   
5               Canada     6      17899              17620     215003   
6              Germany     7      17027              16831     140566   
7                India     8      15005              14841     128763   
8               France     9      13153              12973     130632   
9          South Korea    10      11983              11923     114675   
10               Italy    11      10964              10794     111850   
11               Spain    12       9428               9330     123336   
12           Australia    14       8831      
```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    
    #raise NotImplementedError()
answer_thirteen()
```

Output:
```
Top15
               Country  Rank  Documents  Citable documents  Citations  \
0                China     1     127050             126767     597237   
1        United States     2      96661              94747     792274   
2                Japan     3      30504              30287     223024   
3       United Kingdom     4      20944              20357     206091   
4   Russian Federation     5      18534              18301      34266   
5               Canada     6      17899              17620     215003   
6              Germany     7      17027              16831     140566   
7                India     8      15005              14841     128763   
8               France     9      13153              12973     130632   
9          South Korea    10      11983              11923     114675   
10               Italy    11      10964              10794     111850   
11               Spain    12       9428               9330     123336   
12           Australia    14       8831      
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```